```python
# Initialize Otter
import otter
grader = otter.Notebook("lab03-control.ipynb")
```

# Coding Lab 3: Print, `None`, and Control

## Tuskegee Scholars Summer 2023

Goals of Lab 3:
* Differentiate betwen the `print()` function and the `return` statement
* Understand the properties of `None`
* Build functions with `if-elif-else` conditional statements

While you will not submit these labs, we hope that you can work through the "required" parts of each lab. This will help build your foundation for Data 6, which starts in Week 3 of this program.

You will just be working through this single lab today, which has 6 questions. There are several bonus questions; feel free to work through these once you've finished the main lab.

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 0: Find a Partner!

Labs are best worked on in pairs (or groups of three)! First, find a **new** partner (or two) that you'd like to work with on this lab. Share the following things:

* Introduce yourself (name, major, year, etc.).
* What's your current favorite song?
* What did you have for dinner? Was it good?
* How was your summer so far?

Now that you have introductions out of the way, let's work on the lab together!

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 1: Warmup/Review

Let's review functions!

---

# Question 1: Temperature Converter

How can we communicate the highs and lows of climate change with people who use a different temperature unit? 
Let's make a temperature converter based on the steps here: 
https://www.mathsisfun.com/temperature-conversion.html

Create a function called `celsius_to_fahrenheit` that:
* takes a single argument, the temperature in celsius
* calculates and returns the fahrenheit equivalent

Similarly, create another function called `fahrenheit_to_celsius` that:
* takes a single argument, the temperature in fahrenheit
* calculates and returns the celsius equivalent


Your implementations should correctly evaluate the following:
```
celsius_to_fahrenheit(0)   # 32
celsius_to_fahrenheit(100) # 212
fahrenheit_to_celsius(32)  # 0
fahrenheit_to_celsius(212) # 100
```

```python
# This def statement may be incomplete...
def celsius_to_farenheit():
    ...
    
# This def statement may be incomplete...
def farenheit_to_celsius():
    ...


# edit the below lines as needed to test your functions
print("0 degrees Celsius is", celsius_to_farenheit(0), "degrees Farenheit")
print("0 degrees Farenheit is", farenheit_to_celsius(0), "degrees Celsius")
```

```python
grader.check("q1")
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: `print()`

The `print()` function is used to display values in Jupyter. This can lead to some confusion between print and **output**, because evaluating the last expression in a Jupyter code cell will **output** the resulting value.

However, this is different from **displaying** values directly from the cell (which is what `print()` does).

```python
# printing/displaying
print("Hello World!")
```

```python
# vs outputting
"Hello World!"
```

Another thing to remember is that 'print' is a function itself, so it also has a return value. The return value of 'print' is 'None'.

```python
print(print("Hello World!"))
```

However, note that the following cell only **prints**; it does **not** output.

```python
output_of_print = print("Hello World!")
output_of_print
```

This is because Jupyter code cells are programmed not to output `None` even if it is the last evaluated value in the cell. Check it out:

```python
# this does not output anything
None
```

The below cell doesn't output anything either; it only displays:

```python
print("no output! only print!")
```

By contrast,

```python
2
```

### `return` vs `print`

Let's return to a familiar function:

```python
def square(x):
    return x * x

square(4)
```

Recall that when Python executes a `return` statement, the function terminates immediately. If Python reaches the end of the function body without executing a `return` statement, it will automatically return `None`.

In contrast, the `print()` function is used to **display** values in the notebook. This can lead to some confusion between `print()` and `return` because calling a function in a cell will often **output** the function's return value.

However, unlike a `return` statement, when Python evaluates a `print()` expression, the function does not terminate immediately.

```python
def what_prints():
    print('Hello World!')
    return 'Exiting this function.'
    print('We are awesome!')
```

```python
what_prints()
```

Notice also that print will display text **without the quotes**, but return will preserve the quotes (because it's the output value).

Finally, remember that `print()` is a function itself, so it also has a return value. The return value of `print()` is `None`.

<br/><br/>

---

# Question 2: What Would Python *Display*? Part 1

From [CS 88, Discussion 01](https://cs88-website.github.io/sp22/disc/disc01.pdf).

**Before** running the following cells, consider what (if anything) Python will print, (i.e., display--not output). Then, run the cells. You can use the keyboard shortcut `Shift + <Enter>` to run and step to the next cell.


This question is designed as a discussion question, so talk through your ideas with your partner(s)!

```python
print('x')
```

```python
y = print('x')
```

```python
print(y)
```

```python
from operator import add
print(add(4, 2), print('a'))
```

<br/><br/>

---

# Question 3: What would Python *Display*? Part 2

Remember: `print()` (unlike `return`) does not cause the function to exit!


This question is designed as a discussion question, so talk through your ideas with your partner(s)!

## Question 3a

```python
def ab(c, d):
    if c > 5:
        print(c)
    elif c > 7:
        print(d)
    print('foo')

ab(10, 20)
```

## Question 3b

```python
def bake(cake, make):
    if cake == 0:
        cake = cake + 1
        print(cake)
    if cake == 1:
        print(make)
    else:
        return cake
    return make
```

```python
bake(0, 29)
```

```python
bake(1, "mashed potatoes")
```


<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 3: Conditional Statements

Conditional statements let programs execute different lines of code depending on certain
conditions. Let’s review the if- elif-else syntax:

```
if <conditional expression>:
    <suite of statements>
elif <conditional expression>:
    <suite of statements>
else:
    <suite of statements>
```

Recall the following points:
* The `else` and `elif` clauses are optional, and you can have any number of elif clause.
* A **conditional expression** is a expression that evaluates to either a true value (`True`, a non-zero integer, etc.) or a false value (`False`, `0`, `None`, etc.).
* Only the suite that is indented under the first `if`/`elif` that has a conditional expression that evaluates to True will be executed.
* If none of the conditional expressions are `True`, then the else suite is executed. There can only be one `else` clause in a conditional statement!

**Tip**: You may be tempted to write code like this:

```
if x > 3:
    return True
else:
    return False
````

This can be written more concisely as:

```
return x > 3
```

Boolean expressions and conditional statements take practice! Whenever you write conditional statements that involve returning or binding boolean values, see if you can rewrite your code.


<br/>

---

# Question 4: If True, What Would Python Do? (ITWWPD)

From [CS 88 Lab 01](https://cs88-website.github.io/sp22/lab/lab01/), [CS 61A Lab 01](https://cs61a.org/lab/lab01/#error-messages) and [CS 61A Discussion 01](https://cs61a.org/disc/disc01/).


**Before** running the following cells, consider what Python would do. Then, run the cells. You can use the keyboard shortcut `Shift + <Enter>` to run and step to the next cell.


This question is designed as a discussion question, so talk through your ideas with your partner(s)!



## Question 4a

```python
def xk(c, d):
    if c == 4:
        return 6
    elif d >= 4:
        return 6 + 7 + c
    else:
        return 25
    return 12
```

```python
xk(10, 10)
```

```python
xk(10, 6)
```

```python
xk(4, 6)
```

```python
xk(0, 0)
```

<br/>


## Question 4b

```python
def how_big(x):
    if x < 0:
        print('negative')
    if x > 10:
        print('huge')
    elif x > 5:
        print('big')
    elif x > 0:
        print('small')
    else:
        print('nothing')
```

```python
how_big(7)
```

```python
how_big(12)
```

```python
how_big(1)
```

```python
how_big(-1)
```

<br/><br/>

---

# Question 5: `if` vs `elif`

In this question, we will explore the difference between `if` and `elif`.

All parts of this question are designed as discussion exercises, so talk through your ideas with your partner(s)!

## Question 5a
What is the result of evaluating the following code?


```python
def special_case():
    x = 10
    if x > 0:
        x += 2
    elif x < 13:
        x += 3
    elif x % 2 == 1:
        x += 4
    return x

special_case()
```

_Write your answer in this Markdown cell._

## Question 5b

What about this piece of code?

```python
def just_in_case():
    x = 10
    if x > 0:
        x += 2
    if x < 13:
        x += 3
    if x % 2 == 1:
        x += 4
    return x

just_in_case()
```

_Replace this cell with your answer._

_Write your answer in this Markdown cell._

## Question 5c

What about this one?

```python
def case_in_point():
    x = 10
    if x > 0:
        return x + 2
    if x < 13:
        return x + 3
    if x % 2 == 1:
        return x + 4
    return x

case_in_point()
```

_Write your answer in this Markdown cell._

## Question 5d

Which of the above code snippets result in the same output, and why? Based on your findings, when do you think using a series of `if` statements has the same effect as using both `if` and `elif` cases?

_Write your answer in this Markdown cell._


<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 4: Conditional Statements vs Boolean Expressions

For the last part of this lab, you will implement a function that returns a boolean value and compare two different approaches: (1) a conditional statement or (2) a one-line boolean expression.

# Question 6: Jacket? No Jacket?

From [CS 88 Discussion 01](https://cs88-website.github.io/sp22/disc/disc01.pdf).

Tommy will only wear a jacket outside if it is below 60 degrees or it is raining.
Write a function that takes in the current temperature and a boolean value telling if it
is raining and returns `True` if Tommy will wear a jacket and `False` otherwise.

In other words, the function should correctly evaluate the following:
```
wears_jacket(90, False)  # False
wears_jacket(40, False)  # True
wears_jacket(100, True)  # True
```

## Question 6a
First, try solving this problem using an `if-elif-else` statement.

```python
def wears_jacket(temp, raining):
    ...

# edit the below line as needed for testing
wears_jacket(90, False)
```

```python
grader.check("q6a")
```

<br/>

## Question 6b
Note that we’ll either return `True` or `False` based on a single condition, whose truthiness value will also be either `True` or `False`. 

Knowing this, try to implement this using a single line in the function `wears_jacket_single_line` below.

```python
def wears_jacket_single_line(temp, raining):
    ...

# edit the below line as needed for testing
wears_jacket_single_line(90, False)
```

```python
grader.check("q6b")
```


<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Bonus/More Practice

From [CS 88 Discussion 01](https://cs88-website.github.io/sp22/disc/disc01.pdf), [CS 88 Lab 01](https://cs88-website.github.io/sp22/lab/lab01/), [CS 61A Lab 01](https://cs61a.org/lab/lab01/#error-messages), and [CS 61A Lecture 03](https://inst.eecs.berkeley.edu/~cs61a/sp22/assets/slides/03-Control.html#/32).

# Question 7: Greater Number

Write a function named `greater_num` that:
* takes 2 arguments, both numbers.
* returns whichever number is the greater (higher) number.

Use conditional statements to implement the function.

Your implementation should correctly evaluate the following:

```
greater_num(45, 10)   # 45
greater_num(-1, 30)   # 30
greater_num(20, 20)   # 20
```

```python
def greater_num(num1, num2):
    ...

# edit the below line as needed for testing
greater_num(45, 30)
```

```python
grader.check("q7")
```

# Question 8: Hello, Worldwide!

Write a function named `hello_world` that:
* takes 1 argument, a language code (e.g. `"es"`, `"pt"`, `"en"`)
* returns `"Hello, World"` for the given language
It should return English if an invalid language code is specified.

Your implementation should correctly evaluate the following:

```
hello_world("en")    # 'Hello, World'
hello_world("es")    # 'Hola, Mundo'
hello_world("pt")    # 'Olá, Mundo'
```

We encourage you to expand this function by adding your favorite language!


```python
def hello_world(language_code):
    ...

# edit the below line as needed for testing
hello_world('en')
```

```python
grader.check("q8")
```

# Question 9: Assign Grading Bins

Write a function named `assign_grade` that:
* takes 1 argument, an integer representing the score.
* returns a grade for the score, either `"A"` (90-100), `"B"` (80-89), `"C"` (70-79), `"D"` (65-69), or `"F"` (0-64).

Call that function for a few different scores and log the result to make sure it works.

Your implementation should correctly evaluate the following:
```
assign_grade(95)   # 'A'
assign_grade(90)   # 'A'
assign_grade(85)   # 'B'
assign_grade(80)   # 'B'
assign_grade(77)   # 'C'
assign_grade(70)   # 'C'
assign_grade(68)   # 'D'
assign_grade(65)   # 'D'
assign_grade(64)   # 'F'
assign_grade(0)    # 'F'
```

```python
def assign_grade(score):
    ...

# edit the below line as needed to test your function
assign_grade(95)
```

```python
grader.check("q9")
```

<br/><br/>

---

# Question 10: Classroom Management

To handle discussion section overflow, Matt and other TAs may direct students to a more empty section that is happening at the same time.

Write a function that takes in the number of students in two sections and returns (does not print!) a string about what to do if either section exceeds 30 students.

**Hint**: You can do `str(<number>)+ <string>` to concatenate a number and a string.

Your implementation should correctly evaluate the following:

```
handle_overflow(27, 15)     # "No overflow"
handle_overflow(35, 29)     # "Move to Section 2: 1"
handle_overflow(20, 32)     # "Move to Section 1: 10"
handle_overflow(35, 30)     # "No space left in either section"
```


```python
def handle_overflow(s1, s2):
    if s1 <= 30 and s2 <= 30:
        return "No overflow"
    elif s1 > 30 and s2 < 30:
        return "Move to Section 2: " + str(30 - s2)
    elif s1 < 30 and s2 > 30:
        return "Move to Section 1: " + str(30 - s1)
    else:
        return "No space left in either direction"

# edit the below line as needed to test your function
handle_overflow(35, 30)
```

Output:
```
'No space left in either direction'
```

```python
grader.check("q6")
```

<br/><br/>

---

# [Challenge] Question 11: If "Function" vs If Statement
Now that we've learned about how if statements work, let's see if we can write a function that behaves the same as an if statement.

```python
# just run this cell
def if_function(condition, true_result, false_result):
    """Return true_result if condition is a true value, and
    false_result otherwise.
    """
    if condition:
        return true_result
    else:
        return false_result
```

We'd expect the function to run as follows:

```python
if_function(True, 2, 3)                    # 2
```

```python
if_function(False, 2, 3)                   # 3
```

```python
if_function(3==2, 'equal', 'not equal')    # 'not equal'
```

```python
if_function(3>2, 'bigger', 'smaller')      # 'bigger'
```

Despite what's written above, this function actually does not always do the same thing as an `if` statement!

<br/>

---

We want to find a case where this `if_function` will behave differently from an `if` statement. To do so, implement the following functions:
* `cond`: Calling `cond` should act as the `if` condition.
* `true_func`: Calling `true_func` should represent the result of the truthy case.
* `false_func`: Calling `false_func` should represent the result of the falsey case.

so that `with_if_function` does not behave the same as `with_if_statement` as follows:
* When `with_if_statement` is called, we print out `UC Berkeley`.
* When `with_if_function` is called, we print out both `Welcome to` and `UC Berkeley` on separate lines.

In other words, when you implement the below functions, you should observe the following output:

```
>>>> testing with_if_statement()
UC Berkeley
None

>>>> testing with_if_function()
Welcome to
UC Berkeley
None
```

Implement `cond`, `true_func`, and `false_func` below.

**Hint**: If you are having a hard time identifying how `with_if_statement` and `with_if_function` would differ in behavior, consider the rules of evaluation for `if` statements and call expressions.

```python
# repeated for your reference
def if_function(condition, true_result, false_result):
    if condition:
        return true_result
    else:
        return false_result

def with_if_statement():
    if cond():
        return true_func()
    else:
        return false_func()

def with_if_function():
    return if_function(cond(), true_func(), false_func())

def cond():
    ...

def true_func():
    ...

def false_func():
    ...
    
    
##### do not change anything below this line ######
print(">>>> testing with_if_statement()")
result = with_if_statement()
print(result)
print() # prints a blank line
print(">>>> testing with_if_function()")
result = with_if_function()
print(result)
```

## Done!

That's it! Please feel free to ask any questions!!!

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 2)
```
>>> celsius_to_farenheit(0) == 32
True
```
```
>>> celsius_to_farenheit(100) == 212
True
```
```
>>> farenheit_to_celsius(32) == 0
True
```
```
>>> farenheit_to_celsius(212) == 100
True
```


Tests for q6 (points: 2)
```
>>> handle_overflow(27, 15) == "No overflow"
True
```
```
>>> handle_overflow(35, 29) == "Move to Section 2: 1"
True
```
```
>>> handle_overflow(20, 32) == "Move to Section 1: 10"
True
```
```
>>> handle_overflow(35, 30) == "No space left in either section"
True
```


Tests for q6a (points: 2)
```
>>> wears_jacket(90, False) == False
True
```
```
>>> wears_jacket(40, False) == True
True
```
```
>>> wears_jacket(100, True) == True
True
```


Tests for q6b (points: 2)
```
>>> wears_jacket_single_line(90, False) == False
True
```
```
>>> wears_jacket_single_line(40, False) == True
True
```
```
>>> wears_jacket_single_line(100, True) == True
True
```


Tests for q7 (points: 2)
```
>>> greater_num(45, 10) == 45
True
```
```
>>> greater_num(-1, 30) == 30
True
```
```
>>> greater_num(20, 20) == 20
True
```


Tests for q8 (points: 2)
```
>>> hello_world("en") == 'Hello, World'
True
```
```
>>> hello_world("es") == 'Hola, Mundo'
True
```
```
>>> hello_world("pt") == 'Olá, Mundo'
True
```


Tests for q9 (points: 2)
```
>>> assign_grade(95)   # 'A'
'A'
```
```
>>> assign_grade(90)   # 'A'
'A'
```
```
>>> assign_grade(85)   # 'B'
'B'
```
```
>>> assign_grade(80)   # 'B'
'B'
```
```
>>> assign_grade(77)   # 'C'
'C'
```
```
>>> assign_grade(70)   # 'C'
'C'
```
```
>>> assign_grade(68)   # 'D'
'D'
```
```
>>> assign_grade(65)   # 'D'
'D'
```
```
>>> assign_grade(64)   # 'F'
'F'
```
```
>>> assign_grade(0)    # 'F'
'F'
```
