```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02-functions.ipynb")
```

# Coding Bootcamp Lab 2: Functions

## Tuskegee Scholars Summer 2023

Material is from  [CS 88](https://cs88-website.github.io/sp22/) and [CS 61A](https://cs61a.org/).

Goals of Lab 2:
* Write functions
* Trace through environment diagrams involving functions
* Write boolean functions (more practice with boolean expressions)

If you haven't paired up/teamed up with anyone yet, we strongly encourage you to for this lab! This lab is longer than the other labs and is intentionally written more like "puzzles." So grab a friend, your computer, a piece of paper, and a pen, and let's get started :-)

P.S. -- Feel free to ask your partner what they had for lunch. Maybe you'll discover something new...


<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 1: Functions, Autograders, and Environment Diagrams

If we want to execute a series of statements over and over, we can abstract them away into a **function** to avoid repeating code.

For example, let's say we want to know the results of multiplying the numbers 1-3 by 3 and then adding 2 to it. Here's one way to do it:

```python
1 * 3 + 2
```

```python
2 * 3 + 2
```

```python
3 * 3 + 2
```


If we wanted to do this with a larger set of numbers, that'd be a lot of repeated code! Let's write a function to capture this operation given any input number.


```python
def foo(x):
    return x * 3 + 2
```


This function, called foo, takes in a single **argument** and will **return** the result of multiplying that argument by 3 and adding 2.

Now we can **call** this function whenever we want this operation to be done:

```python
foo(1)
```

```python
foo(2)
```

```python
foo(1000)
```

Applying a function to some arguments is done with a **call expression.**


<br/><br/>


## Otter Grader

In this lab, we are introducing a new Jupyter system that you will use extensively in Data 6: **Otter Grader.** Otter grader is an autograder (yes, the similar pronunciation is intentional), meaning that it checks your Python code for correctness.

If you haven’t already, make sure you have run the cell at the top of this notebook to initialize Otter.

With Otter, you can start using the built-in tests to check whether your work is correct. Sometimes, there are multiple tests for a single question, and passing all of them is required to receive credit for the question. Please don't change the contents of the test cells.

Go ahead and attempt the question below. Running the cell directly after it (`grader.check(...)`) will test whether you have implemented `ilovedata()` correctly. If you haven't, this test will tell you the correct answer.


---

# Question 1

Complete the below question `ilovedata()` below such that it returns the string `"I love Data Science!"`.

```python
def ilovedata():
    # BEGIN SOLUTION
    return "I love Data Science!"
    # END SOLUTION

# We've put this line in this cell 
# so that it will print the value you've given to seconds_in_a_decade when you run it.  
# You don't need to change this.
ilovedata()
```

Output:
```
'I love Data Science!'
```

```python
grader.check("q1")
```

<br/><br/>


## More Environment Diagrams

From [CS 61A Discussion 01](https://cs61a.org/disc/disc01/).

Let's come back to the environment diagram. It is a visual way of understanding how Python works with names (i.e., variables) and function calls.

One key idea in environment diagrams is the **frame**. A frame helps us keep track of what variables (i.e., names) have been defined in the current execution environment, and what values they hold. The frame we start off with when executing a program from scratch is what we call the **Global frame**. Later, we'll get into how new frames are created and how they may depend on their parent frame.

<br/>

---

# Question 2: Discussing and Tracing a Program

Here's a short program and its corresponding diagram:

**Python Tutor** [click here](https://pythontutor.com/render.html#code=x%20%3D%203%0A%0Adef%20square%28x%29%3A%0A%20%20%20%20return%20x%20*%20x%0A%20%20%20%20%0Ax%20%3D%20square%282%29%0Ax&cumulative=false&curInstr=0&heapPrimitives=nevernest&mode=display&origin=opt-frontend.js&py=3&rawInputLstJSON=%5B%5D&textReferences=false)

```python
x = 3

def square(x):
    return x * x

x = square(2)
x
```

With your partner and the Python Tutor:
* Trace through the above program.
* What is the final value of `x`?
* Identify the different frames that are created.
* How is the return value of the `square` function visualized in the environment diagram?
* Identify all the values that the name `x` binds to. Compare and contrast. How are they different?


<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: Function Practice

From [CS 61A Lecture 02](https://cs61a.org/assets/slides/02-Functions.html#/39).

To improve coding skills, you will need to *practice, practice, practice.*

Many (if not all) of the programs below can be written with a single line in the body, like so:

```
def some_function(arg):
    return some_expression_with_arg
```

However, sometimes you will find it easier to create *intermediate* assignment statements to break up your work, like so:

```
def some_function(arg):
    intermediate_name = some_expression_with_arg
    another_intermediate_name = some_expression_with_intermediate_name_and_arg
    return some_final_expression
```

If you follow the latter path, challenge yourself to edit your code so that it fits in one line! Overall, *it is easier to fix messy code that works than to start from a completely blank slate.* (Just like writing essays for humanities courses--ideas are hard, and editing is easier.) So just take your first step, and revise it afterwards!

<br/>

Here are some tips for each exercise below:
* First, start from the arguments (do you need any?).
* Then, think about the return value. What type is it?
* Then, think about the logic to meet the problem specification
* Finally, look back over your code to see if it can be simplified.

<br/><br/>

---

# Question 3: Calculate Dog Age

You know how old a dog is in human years, but what about dog years? Calculate it!

Write a function named `calculate_dog_age` that:
* takes 1 argument: a dog's age in human years.
* calculates the dog's age based on the conversion rate of 1 human year to 7 dog years.
* returns the age in dog years.

Your implementation should correctly evaluate the following:

```
calculate_dog_age(1)   # 7
calculate_dog_age(0.5) # 3.5
calculate_dog_age(12)  # 84
```

```python
def calculate_dog_age(years):
    return years * 7 # SOLUTION

calculate_dog_age(1)
```

Output:
```
7
```

```python
grader.check("q3")
```

<br/><br/>

---

# Question 4: Fortune Teller

Why pay a fortune teller when you can just program your fortune yourself?

Write a function named `tell_fortune` that:
* takes 3 arguments: job title, partner's name, geographic location.
* returns a fortune of the form: 'You will be a {job_title} in {location} living with {partner}.'
**Tip**: If your line of code is getting long, you can split up the line
by using `\` (the backslash character, located above the <Enter> key on a U.S. keyboard) between values in an expression. For example,
    
    "this super long string yaddayadda blah blah blah"
    
turns into (using the concatenator operator `+`)
    
    "this super long string " + \
    "yadda yadda blah blah blah"
    
Note the space at the end of `string `! Why do we keep it?
    
<br/>
    

Your implementation should correctly evaluate the following:

```
tell_fortune('football player', 'Spain', 'Shakira') # 'You will be a football player in Spain living with Shakira.'
tell_fortune('farmer', 'Kansas', 'C3PO') # 'You will be a farmer in Kansas living with C3PO.'
tell_fortune('Elvis Impersonator', 'Russia', 'Karl the Fog') # 'You will be a Elvis Impersonator in Russia living with Karl the Fog.'
```


```python
def tell_fortune(job_title, location, partner):
    return "You will be a " + job_title + " in " + location + " living with " + partner + "." # SOLUTION

# edit the below line as needed to test your function
tell_fortune('football player', 'Spain', 'Shakira') 
```

Output:
```
'You will be a football player in Spain living with Shakira.'
```

```python
grader.check("q4")
```

<br/><br/>

---

# Question 5: Lifetime Supply

## Question 5a

Ever wonder how much a "lifetime supply" of your favorite snack is? Wonder no more!

Write a function named `calculate_supply` that:
* takes 2 arguments: age, amount per day.
* calculates the amount consumed for rest of the life (based on a constant max age of 81).*
* returns the result.

**Hint**: Assume each year has exactly 365 days (no leap years).

\*Side note unrelated to the coding question: This max age is higher than the average life expectancy in the U.S; check out the [World bank](https://data.worldbank.org/indicator/SP.DYN.LE00.IN) for average life expectancy per country as of 2020.



Your implementation should correctly evaluate the following:
```
calculate_supply(80, 1) # 365
calculate_supply(80, 2) # 730
calculate_supply(36, 3) # 49275
```

```python
""" # BEGIN PROMPT
# This def statement may be incomplete...
def calculate_supply():
    return ...
""" # END PROMPT

# BEGIN SOLUTION NO PROMPT
def calculate_supply(age, amount_per_day):
    return (81 - age)*amount_per_day*365
# END SOLUTION

# edit the below line as needed to test your function
calculate_supply(80, 1)
```

Output:
```
365
```

```python
grader.check("q5a")
```

---

## Question 5b

Repeat Question 5a, but accept floating point values for amount per day, and round the result up to an integer.

**Hint**: Should you use `/`, `//`, `round()`, `int()`, or something else?


Your implementation should correctly evaluate the following:
```
calculate_supply_bonus(80, 1) # 365
calculate_supply_bonus(80, 2) # 730
calculate_supply_bonus(36, 3) # 49275
calculate_supply(37, 2.17) # 34850
```

```python
""" # BEGIN PROMPT
# This def statement may be incomplete...
def calculate_supply_bonus():
    return ...
""" # END PROMPT

# BEGIN SOLUTION NO PROMPT
def calculate_supply_bonus(age, amount_per_day):
    return round((81 - age)*amount_per_day*365)
# END SOLUTION

# edit the below line as needed to test your function
calculate_supply_bonus(37, 2.17)
```

Output:
```
34850
```

```python
grader.check("q5b")
```


<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 3: Still More Functions, but Now With Booleans

From [CS 61A Lecture 03](https://cs61a.org/assets/slides/03-Control.html#/19) and [CS 88 Lab 01](https://cs88-website.github.io/sp22/lab/lab01/).

# Question 6: Fix the Bug

The `both_positive()` function is supposed to return `True` if both `x` and `y` are positive (and `False` otherwise), but it doesn't work!

Figure out what is wrong and fix the bugs. A correct implementation should evaluate the following:
```
both_positive(-1, 1)   # False
both_positive(1, 1)    # True
```

```python
def both_positive(x, y):
    # BEGIN SOLUTION NO PROMPT
    return (x > 0) and (y > 0)
    # END SOLUTION
    """ # BEGIN PROMPT
    return x and y > 0 # You can replace this line!
    """ # END PROMPT
    

# edit the below line as needed to test your function
both_positive(-1, 1)
```

Output:
```
False
```

```python
grader.check("q6")
```

<br/><br/>

---

# Question 7: Curly Hair?

Implement the function `has_curly_hair()` below, which should only return `True` if both the father allele and the mother allele are `"C"` (and return `False` otherwise).

**Use only boolean logical and comparator operators**!

Your implementation should correctly evaluate the following:

```
has_curly_hair("C", "c")  # False
has_curly_hair("c", "C")  # False
has_curly_hair("s", "s")  # False
has_curly_hair("C", "C")  # True
```

```python
def has_curly_hair(father_allele, mother_allele):
    # BEGIN SOLUTION NO PROMPT
    return (father_allele == 'C') and (mother_allele == 'C')
    # END SOLUTION
    """ # BEGIN PROMPT
    return ...
    """ # END PROMPT
    
# edit the below line as needed to test your function
has_curly_hair("C", "c")
```

Output:
```
False
```

```python
grader.check("q7")
```

## Done!

That's it! There's nowhere for you to submit this, as labs are not assignments. However, please ask any questions you have with this notebook in lab.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 2)
```
>>> ilovedata() == "I love Data Science!"
True
```
```
>>> len(ilovedata()) == len("I love Data Science!")
True
```


Tests for q3 (points: 2)
```
>>> calculate_dog_age(1) == 7
True
```
```
>>> calculate_dog_age(0.5) == 3.5
True
```
```
>>> calculate_dog_age(12) == 84
True
```


Tests for q4 (points: 2)
```
>>> tell_fortune('football player', 'Spain', 'Shakira') == 'You will be a football player in Spain living with Shakira.'
True
```
```
>>> tell_fortune('farmer', 'Kansas', 'C3PO') == 'You will be a farmer in Kansas living with C3PO.'
True
```
```
>>> tell_fortune('Elvis Impersonator', 'Russia', 'Karl the Fog') == 'You will be a Elvis Impersonator in Russia living with Karl the Fog.'
True
```


Tests for q5a (points: 2)
```
>>> calculate_supply(80, 1) == 365
True
```
```
>>> calculate_supply(80, 2) == 730
True
```
```
>>> calculate_supply(36, 3) == 49275
True
```


Tests for q5b (points: 2)
```
>>> calculate_supply_bonus(80, 1) == 365
True
```
```
>>> calculate_supply_bonus(80, 2) == 730
True
```
```
>>> calculate_supply_bonus(36, 3) == 49275
True
```
```
>>> calculate_supply_bonus(37, 2.17) == 34850
True
```


Tests for q6 (points: 2)
```
>>> both_positive(-1, 1) == False
True
```
```
>>> both_positive(1, 1) == True
True
```


Tests for q7 (points: 2)
```
>>> has_curly_hair("C", "c") == False
True
```
```
>>> has_curly_hair("c", "C") == False
True
```
```
>>> has_curly_hair("s", "s") == False
True
```
```
>>> has_curly_hair("C", "C") == True
True
```
